---
title: "Week 2: Exploratory & Overlap analysis"
tutorial:      # where learnr stores the progress
  id: "FG2"    # folder name
  version: 0.1 # subfolder name
output: 
  learnr::tutorial:
    fig_caption: no
    progressive: true
    allow_skip: true
    css: ["css/custom.css", "css/print2pdf.css"]
runtime: shiny_prerendered
description: >
  Multi line text can be put here
  for your convenience.
---


```{r setup, include=FALSE}
library(learnr)     # 0.10.1.9006 (github)
library(gradethis)  # 0.1.0.9004  (github)
library(testthat)   # 3.0.0 
library(tidyverse)  # 1.3.0
library(learnr.proto)

library(IRanges)
library(GenomicRanges)
library(plyranges)
library(rtracklayer)
library(VennDiagram)
library(GenomicFeatures)
library(ComplexHeatmap)
library(genomation)
library(ChIPpeakAnno)
library(ggthemes)
library(org.Hs.eg.db)


# load the txdb package which holds transcript-based gene models of hg38 genome  
library(TxDb.Hsapiens.UCSC.hg38.knownGene)
txdb <- TxDb.Hsapiens.UCSC.hg38.knownGene 
seqlevels(txdb) <- "chr19"

## installing additional packages by:
# BiocManager::install('genomation')
# BiocManager::install("ggthemes")
# BiocManager::install("ComplexHeatmap")  # for upset plot with GRanges
# BiocManager::install("org.Hs.eg.db")


# configuration
knitr::opts_chunk$set(echo = FALSE, message = F)
gradethis::gradethis_setup()

# data accessible to all exercises
rdata = system.file("extdata", "week2.Rdata", package = "learnr.proto")
load(rdata)
monocytes_all <- unlist(monocytes_list)
monocytes_all$chip <- names(monocytes_all)
monocytes_list <- GRangesList(monocytes_list, compress = F)

# chr19_windows_1bp <- tileGenome(seqinfo(txdb), tilewidth = 10, cut.last.tile.in.chrom =T)
promoters_chr19 <- unique(promoters(genes(txdb), upstream=2000, downstream=0)) 
tss_chr19  <- resize(promoters_chr19, width = 1, fix = "end")
genes_chr19 <- unique(genes(txdb))

# gene quantification chromosome 19  
rdsfile <- system.file("extdata", "week2", "prepared_rds", "blueprint_c000s5_gene_quantification_chr19.rds", package = "learnr.proto")
quantification_chr19 <- readRDS(rdsfile)

### not used anymore 2020-12-14 ###
# rdata <- system.file("extdata", "week2", "prepared_rds", "txdb_annotation_chr19_genes_tx_pr_tss_granges.RData", package = "learnr.proto")
# load(rdata)
# identify TSS (in this case 1 per gene even though we know that genes have multiple TSSs)
# tss <- resize(genes, width = 1, fix = "start")

# consider 1kb upstream and 200bp downstream the TSS as promoter
# promoters <- resize(tss, width = 1000, fix = "end")
# promoters <- resize(promoters, width = 1200, fix = "start")

# windows_1bp <- tileGenome(seqinfo(txdb), tilewidth = 10, cut.last.tile.in.chrom =T)
# saveRDS(annotation_chr19,"~/learnr.proto/renv/library/R-3.6/x86_64-pc-linux-gnu/learnr.proto/extdata/week2/prepared_rds/genomation_chr19_background_annotation.rds")

#  genomation_chr19 <- system.file("extdata", "week2", "prepared_rds","genomation_chr19_background_annotation.rds", package = "learnr.proto")

# import the gencode features for annotation with genomation (exercise 8)
# gencode_chr19 <- system.file("extdata", "week2", "encode", "gencodev32.chr19.bed", package = "learnr.proto")
# gencode_chr19 <- readTranscriptFeatures(gencode_chr19, up.flank = 2000, down.flank = 200) 



```

```{r timeout1, context="server"}
isolate({
  # obtain the session id
  # source 1: https://shiny.rstudio.com/reference/shiny/latest/session.html
  # source 2: https://stackoverflow.com/questions/18900955/get-environment-identifier-in-r
  session_id <- sub('<environment: (.*)>', '\\1', capture.output(session$userData))
  
  # # send the session ID to the log (indicates the tutorial is loaded)
  # write(paste0("Session ID: ", session_id), stderr())
  
  # send the session ID to the javascript chunk
  session$sendCustomMessage("session_id", session_id)
})

# stop the tutorial when "input[[session_id]]" is updated
observeEvent(input[[session_id]], ignoreNULL=T, {
  write(paste0("\nTutorial terminated due to inactivity.\nRestart to continue where you left off!\n"), stderr())
  stopApp()
})
```

```{js timeout2}
// This chunk stops the tutorial if it has been idle for too long

// Situation: We need to stop inactive tutorials to prevent the server from grinding to a halt.
// Problem: refreshing the tutorial and closing the browser send the same signal to shiny.
// since we cannot distinguish the two, this signal cannot be used to stop the tutorial.
// solution: Use a timeout system. 
// Method: this chunk will update shiny variable "input[[session_id]]" 
//   when the session has been running idle for longer than "timeoutSec" seconds.
// This signal is then be used by Shiny (server-side) to stop the session.

// source1: https://community.rstudio.com/t/keeping-track-of-idle-time-during-app-usage/1735
// source2: https://bookdown.org/yihui/rmarkdown/language-engines.html#javascript-and-css
$(function() {
  var timeoutSec = 5*60;
  var idleTimer;

  // receive this session's ID
  Shiny.addCustomMessageHandler("session_id", function(s_id) {
    session_id = s_id;  // assigns the variable globally
  });

  // assign session ID as reactive variable "input[[session_id]]" 
  function onTimeout() {
    alert("Tutorial stopped due to inactivity.\nRestart to continue where you left off!")
    Shiny.setInputValue(session_id, "TRUE");
  }

  function startIdleTimer() {
    if (idleTimer) clearTimeout(idleTimer);
    idleTimer = setTimeout(onTimeout, timeoutSec * 1000);
  }
  
  $(document).on('shiny:message shiny:inputchanged', startIdleTimer);

})();
```


## 2.1 Introduction
<br>
In week 1 we discussed the major steps involved in obtaining genomics data, from experiment to raw data to normalized signal and peaks. You examined histone PTM ChIP-, DNaseI-, and RNA-seq data of monocytes in the UCSC genome browser and searched for regions with increased signal, *peaks*. You looked at the position of these peaks in the chromosome and with respect to genes as well as the co-occurrence of different marks.   
<br>
These observations showed examples of marks associated with a particular functional elements (e.g. active enhancer, active promoter, silenced domains). Manual annotation of each peak region would be an overkill and dangerously subjective. 
<br>
**This week and next week you will perform computational analyses to annotate peaks and answer common questions including:** "*How many peaks do I have?*", "*Is this mark statistically enriched at promoters?*", "*Which and how often do marks co-occur?*","*The expression of which gene could be affected by this mark?*",  and "*What is the signal of the mark around a particular element of interest like the TSS?*"   
<br>

### 2.1.1 Learning Objectives  

> At the end of week 2 you are able to:  
> 
1. Import ChIP-seq peaks into a GRanges object in r.  
2. Perform exploratory data analysis on GRanges objects with ChIP-seq peaks using plyranges and ggplot. 
3. Detect and count overlap between two GRanges objects.    
4. Plot this overlap in a venndiagram.  
5. Statistically test for enrichment of histone marks in a particular genomic region.  
6. Visually compare gene expression of genes with different histone marks at their promoters.  
7. Summarize the genome-wide overlap among multiple marks in a upset plot.  
- These refer to global learning objectives #4-#7.  

### Exercises
In this tutorial you will have multiple-choice and checkbox questions like in **fg1** but also complete-the-code exercises. The latter can be evaluated on the computed outcome (e.g a computed p-value), in other cases, the code itself is checked. 
<br>
For short lines of code, you have to decide which function or object to use and how to write this in code. For some longer lines of codes, we left blanks as "___" that indicate where in the code you should add a function or object name.  
<br>

- Hit `Start Over` to remove any adjustments you made and start again.   
- Hit `Hint` (when provided) for clues.  
- Hit `Run Code` to preview your code and the output.  
- Hit `Submit Answer` (if present) to submit the code. As 'submitting' not always shows the code output, it is advised to use `Run Code` first.  

<br>
```{r exemplemc}
quiz(caption = "Example questions:",
     question("What number is the letter A in the English alphabet?",
        answer("8"),
        answer("14"),
        answer("1", correct = TRUE),
        answer("23")
     ),
     question("Where are you right now? (select ALL that apply)",
          answer("Planet Earth", correct = TRUE),
          answer("Pluto"),
          answer("At a computing device", correct = TRUE),
          answer("In the Milky Way", correct = TRUE),
          incorrect = "Incorrect. You're on Earth, in the Milky Way, at a computer.")
)
```

**Example coding exercise:**  

> **Example exercise**: Print the head of object `monocytes_h3k4me3`.   

```{r examplecode, exercise = T}
# print the head of `monocytes_h3k4me3`

```

<div id="examplecode-hint">
**Hint:** You may want to use the `head(...)` function.  
</div>

```{r examplecode-solution, eval = F}
head(monocytes_h3k4me3)
```

```{r examplecode-check-code, eval = F}
grade_code()
```

<br>  

## 2.2 ChIP-seq peaks in R  {data-progressive=TRUE}  
### 2.2.1 Data formats      
In genomics you often work with **interval** data. Think of peaks, genes, exons, ... any genomic region reported with the **genomic coordinates: chr, start, and end**. As we discussed in week 1, these often come in a tabular format with the basic information about the location and some other information.  
<br>
Common file formats include BED format for peak files and GFF for gene annotations. In practice, peak files come in two flavors, both cohering to the BED format but with small differences in columns 6-7. These are called **narrowPeak** and **broadPeak** files. This is because different data types have different peak shapes and peak-calling takes that into account. If you have narrowPeak files, this means that during peak calling, settings were used that fit ChIP-seq datasets with sharp, narrow enrichment signals. For broadPeak files, peak calling settings were used to detect broad domains of (overall lower) enrichment. We will come back to this difference shortly.      
 

>#### **Background info**: File formats   
<details><summary>What do BED and GFF files hold and look like?</summary>   
**Browser Extensible Data - BED format**:  
>
- Used for peaks, motif locations or other custom intervals.  
- Has 3 required columns: chromosome, start and end position.  
- 6 or 7 optional columns, in case of peaks:  
  4th: peak name  
  5th: peak score (-10log(q-value) * 10, rounded down to integer value))  
  6th: strand to denote orientation (if applicable, otherwise "\*" or "\." if unstranded)  
  7th: signalValue  
  8th: p-value to denote statistical significance, given as -log10(p-value)  
  9th: q-value statistical significance using false discoveray rate, given as -log10(q-value)  
  10th (only for sharp peaks not broad domains): location of peak summit relative to the "start" coordinate.  
>
> Example:  
![](images/wk1_3_file_format_bed.png){width=80%}  
> Further reading on the file formats on the UCSC Genome Browser FAQ page: [BED](http://genome.ucsc.edu/FAQ/FAQformat.html#format1), and the [narrowPeak BED](http://genome.ucsc.edu/FAQ/FAQformat.html#format12) and [broadPeak BED](http://genome.ucsc.edu/FAQ/FAQformat.html#format13).  
> <br>
> **General Feature Format - GFF format**:   
>  
-	Common file format for storing gene annotations not only including genes but also transcripts/splice variants, cDNA sequences, exons, rRNA, ncRNA, etc.  
-	Begins with meta-data in headerlines, starting with #  
-	Records reported in 9 fixed columns  
-	Column 9 can contains various attirbutes (eg gene symbol, the transcript to which the exon belongs).  
- Downloaded for example from [Ensembl](https://www.ensembl.org/Homo_sapiens/Info/Index).  
>
> Example:   
![](images/wk1_3_file_format_gff.png){width=80%}      
> Further reading on the gff file formats: [UCSC Genome Browser](http://genome.ucsc.edu/FAQ/FAQformat.html#format3), or [ENSEMBL](https://m.ensembl.org/info/website/upload/gff3.html).  
> <br>


### 2.2.2 Peaks as data.frame     
<br>
Peak data can be imported into a data.frame with:  
```{r exampleimport, eval =F, echo =T}
object_name <- read.table("location/of/peak_file.bed")
```
<br>
The data.frame data structure is, however, not the most efficient way to work with interval data. E.g. a simple manipulation such as shifting all reported intervals 2 bp to the right, requires you to manipulate the "start" and "end" columns at the same time.  
<br>
Interval data such as genomic peaks can be more efficiently handled with the IRanges package which works with a data structure especially developed for **ranges of integers**: IRanges objects.  

### 2.2.3 Introducing IRanges and GRanges    
To construct an IRanges object you need to define at least two of the following three values:  

1. a starting coordinate  
2. a finishing coordinate  
3. the width of the interval.  

<br>
The GRanges objects of the GenomicRanges package are very similar but require a additional **sequence name** (in other words a chromosome) for every interval and an optional **strand** column.  
<br> 
Here we create GRanges object called `gr` with 3 interals on chr1:  
```{r create gr, echo = T, eval = T}
# create `gr`
gr_chr1 <-  GRanges(seqnames = c("chr1"), # argument is recycled  
              ranges = IRanges(start=c(5,10,25), width = c(13,17,6)),
              strand = c("-","+", "+"),
              color = c("black", "red", "green"))
# add names to `gr_chr1`
names(gr_chr1) <- paste("interval", LETTERS[1:3], sep = "_")
# print the object:
gr_chr1
```
<br>
When we visualize these intervals we see three 'blocks' along the horizontal axis:
```{r plot gr, echo = F}
plotRanges(gr_chr1, col = gr_chr1$color)
rm(gr_chr1)
```

### 2.2.4 Handy functions for GRanges    
There are several functions to retrieve and set the values in GRanges objects. Here is an image of a GRanges object called `gr` with the respective functions. You will use these functions soon.  

![](images/wk2_2_GRanges_fields_my_gr.png){width=100%}   
<br>
A list of the above functions (and more) for GRanges can also be called with `methods(class = "...")`. This also command works for other objects as well. You just need to figure out the class of the object in question using `class([object_name])`.  
<br>
You can print a summary of the GRanges object with `show(object_name)` and `print(object_name)` (this works for many  other R objects as well).  
<br>

### 2.2.5 Slots in GRanges objects     
Notice that `gr` holds 3 slots:  

1. the genomic coordinates in a **GRanges** slot  
2. 2 columns as **metadata**  (metadata is data about the data)    
3. a genome in the **seqinfo** slot  

The main differences between the metadata and the genomic coordinates are:  

Genomic coordinates |  Metadata columns   
:--|:---  
Printed on the left-hand side of the \|-sign | Printed on the right-hand side of the \|-sign     
Extract using `granges(object_name)` |  Extract as DataFrame with `mcols(object_name)` or `object_name$column_name` for a specific column   
Restricted to variables `seqnames`, `ranges` and `strand` | Almost anything can be stored in the metadata

<br>
Information about the genome is stored in the **seqinfo** part of the object. In the example you see that the peaks are located in human genome 38 which holds 455 sequences in total of which 1 circular (mitochondrial genome). (Sequences that can not be confidently placed on a specific chromosome is stored in chrUn, chr*N*_random and chr*N*_region [explanation on UCSC GB website](http://genome.ucsc.edu/FAQ/FAQdownloads#download10).)   


### 2.2.6 Importing peak files as GRanges     
To parse peak files directly into GRanges you use the `import()` function from the rtracklayer package, for example:    
```{r exampleimport_rtracklayer, eval = F, echo = T}
# load the required library, is usually done once at the start of your analysis.  
library(rtracklayer)

# import monocyte H3K4me3 peak locations
monocytes_h3k4me3 <- import("C000S5H2.ERX547984.H3K4me3.bwa.GRCh38.20150527.chr19.bed", format = "narrowPeak")

# (you could also use rtracklayer::import("C000S5H2.ERX547984.H3K4me3.bwa.GRCh38.20150527.chr19.bed", format = "narrowPeak") and skip loading the library)
```
<br>
As mentioned above (2.2.1), different data types have different peak shapes resulting in "narrowPeak" and "broadPeak" files. H3K4me3 ChIP-seq generally shows sharp(er) enrichment that can reach high enrichment values. For `import()` to function properly, we defined this format in the `format = ` parameter.  
<br>
H3K4me1 ChIP-seq cover broad(er) domains and peak calling took that into account, resulting in a "broadPeak" peak file. 
```{r qimport}
question("Which function imports the H3K4me1 dataset into a GRanges object?",
         answer("import(\"C000S5H2.ERX547981.H3K4me1.bwa.GRCh38.broad.20150527.chr19.bed\", format = \"narrowPeak\")", message = "check the file format."),
         answer("import(\"C000S5H2.ERX547981.H3K4me1.bwa.GRCh38.broad.20150527.chr19.bed\", format = \"BED\")", message = "BED is not the right format."),
         answer("import(\"C000S5H2.ERX547981.H3K4me1.bwa.GRCh38.broad.20150527.chr19.bed\", format = \"broadPeak\")", correct =T),
         answer("read.table(\"C000S5H2.ERX547981.H3K4me1.bwa.GRCh38.broad.20150527.chr19.bed\", format = \"broadPeak\")", message = "read.table parses the data into a data.frame while we asked for a GRanges object."),
         allow_retry = TRUE,
         random_answer_order = TRUE
         )
```
<br>
Let's take a quick look at `monocytes_h3k4me3`.     

### Exercise 1  

> **Exercise 1a**: Print the head of `monocytes_h3k4me3`.  

```{r q1_head_h3k4me3, exercise = TRUE}
# print the head of monocytes_h3k4me3  
# the object "monocytes_h3k4me3" has already been imported for you


```

<div id="q1_head_h3k4me3-hint">
**Hint:** Use the `head(...)` function.  
</div>

```{r q1_head_h3k4me3-solution, eval = F}
head(monocytes_h3k4me3)
```

```{r q1_head_h3k4me3-check}
grade_result(
 pass_if(~identical(.result, head(monocytes_h3k4me3)))
 )
```

```{r quiz_grangesobject}
quiz(caption = "", 
    question("How many metadata columns does the 'monocytes_h3k4me3' object have?",
             answer("2"),
             answer("3"),
             answer("6", correct = T, message = "The original .BED files lacked column headers. The `rtracklayer::import()` function appendended variable names and converted the data to the data type (eg numeric, integer, character etc.) according to the file format we defined."),   
             answer("9"), 
             allow_retry = TRUE),
    question("How many different chromosomes are reported in this file? *Hint*: look in seqinfo part of the object.",
             answer("1", correct = T, message = "We have restricted all peak files to chromosome 19 to limit their file size. That is why 'seqinfo' holds only one sequence"),
             answer("24"),
             answer("455"),
             allow_retry = TRUE)
)
```

### Briefly summarizing  
You have the following information in your object `monocytes_h3k4me3`:  

Column | Section in GRanges object | Variable | Description   
-- | --- | --- | ------  
1. | GRanges | seqnames | chromosome name of peak location.        
2. | GRanges | ranges | start and enc position of peak.         
3. | GRanges | strand | if applicable, orientation of the peak, otherwise \*\.        
4. | Metadata | name | peak name, given by the peak caller.     
5. | Metadata | score | -10log(qvalue) * 10, rounded down to integer value.  
6. | Metadata | signalValue | enrichment signal at the peak summit.  
7. | Metadata | pValue | -log10(p-value), (*e.g.if p-value = 1e-10, this value is 10*), roughly put, the significance of the enrichment.   
8. | Metadata | qValue | -log10(qvalue), the false discovery rate determined by swapping test and control.  
9. | Metadata | peak  | location of peak summit relative to the peak start.  

## 2.3 Exploratory data analysis   
### 2.3.1 Peak counts
A common first step in genomics data analysis is to make look at some basic characteristics of our data by making exploratory summaries and plots.  

### Exercise 2a    

> **Exercise 2a**: Determine the number of H3K4me3 and H3K4me1 peaks in monocytes. 
> 
> * The peak data are stored in the R objects `monocytes_h3k4me3` and `monocytes_h3k4me1`.    

```{r q2a_exploreh3k4me3, exercise = T}
# How many peaks are stored in monocytes_h3k4me3?
monocytes_h3k4me3

```

<div id="q2a_exploreh3k4me3-hint">
**Hint:** You may want to use the `length()` function.  
</div>

```{r q2a_exploreh3k4me3-solution, eval = F}
length(monocytes_h3k4me3)
```

```{r q2a_exploreh3k4me3-check, eval = F}
grade_result(
 pass_if(~identical(.result, length(monocytes_h3k4me3)))
 )
```

```{r q2a_exploreh3k4me1, exercise = T}
# And how many H3K4me1 peaks do you have?


```

```{r q2a_exploreh3k4me1-solution, eval = F}
length(monocytes_h3k4me1)
```

```{r q2a_exploreh3k4me1-check, eval = F}
grade_result(
 pass_if(~identical(.result, length(monocytes_h3k4me1)))
 )
```

### Exercise 2b  

> **Exercise 2b**: And what is the distribution of peak sizes for these marks? 

```{r q2b_widthh3k4me3, exercise=TRUE, exercise.eval = F, error = T}
# What is the distribution of H3K4me3 peak sizes?   
summary(___(monocytes_h3k4me3))
```

<div id="q2b_widthh3k4me3-hint">
**Hint:** You can retrieve peak widths with the `width()` function.  
</div>

```{r q2b_widthh3k4me3-solution, eval = F}
summary(width(monocytes_h3k4me3))
```

```{r q2b_widthh3k4me3-check, eval = F}
grade_result(
  pass_if(~identical(.result, summary(width(monocytes_h3k4me3))))
)
```

```{r q2b_widthh3k4me1, exercise=TRUE, exercise.eval = F, error = T}
# And what is the distribution of H3K4me1 peak sizes?   
monocytes_h3k4me1
```

```{r q2b_widthh3k4me1-solution, eval = F}
summary(width(monocytes_h3k4me1))
```

```{r q2b_widthh3k4me1-check, eval = F}
grade_code(
  pass_if(~identical(.result, summary(width(monocytes_h3k4me1))))
)
```

```{r quiz_peakgranges}
quiz(caption = "", 
    question("What are the median peak sizes for H3K4me3 and H3K4me1 ChIP-seq peaks in the current dataset? (in monocytes, restricted to chromosome19)",
             answer("median H3K4me3 peak = 879.5 bp", correct = T),
             answer("median H3K4me3 peak = 1662 bp"),
             answer("median H3K4me3 peak = 1198.6 bp"),   
             answer("median H3K4me3 peak = 2642 bp"),
             answer("median H3K4me1 peak = 879.5 bp"),
             answer("median H3K4me1 peak = 1662 bp", correct = T),
             answer("median H3K4me1 peak = 1198.6 bp"),   
             answer("median H3K4me1 peak = 2642 bp"), 
             allow_retry = TRUE),
    question("And what are the widths of the *largest* peaks in the current H3K4me3 and H3K27me1 ChIP-seq datasets of monocytes?",
             answer("largest H3K4me3 peak = 8237 bp", message = "8237 bp is the width of the largest H3K27ac peak on chr19" ),
             answer("largest H3K4me3 peak = 12705 bp", correct = T),
             answer("largest H3K4me3 peak = 73112 bp", message = "73112 bp is the width of the largest H3K36me3 peak on chr19" ),   
             answer("largest H3K4me1 peak = 32063 bp", correct = T ),
             answer("largest H3K4me1 peak = 86201 bp", message = "86201 bp is the width of the largest H3K27me3 peak on chr19"),
             answer("largest H3K4me1 peak = 136766 bp", message = "136766 bp is the width of the largest H3K9me3 peak on chr19"),
             allow_retry = TRUE))
```
<br>
Notice the larger median and maximum peak size for H3K4me1 peaks compared to H3K4me3 peaks. This is related to these datasets having relative "broad domain" and "narrow peaks" of enrichment, respectively.  

### 2.3.2 Plotting summaries with ggplot2  
You will use the ggplot2 package to plot the number of peaks for each dataset, to compare the average peak sizes and to plot the total number or fraction of bp's covered by these peaks on chromosome 19.  
<br>

> #### A very brief refresher on ggplot2-ting:   
> With **ggplot2** you generally begin with the `ggplot(data = ..)` function and you add so-called layers with `geom..()` functions. Example geom functions include `geom_point()` for a scatter plot, `geom_histogram()` for a histogram, `geom_density()` for a smoothened version of the histogram, and `geom_bar()` or `geo_col()` for a barplot.  
> <br>
> The **geom functions** take a `mapping = aes(x = ..., y = ..., ...)` argument that defines how variables in the dataset are mapped to visual properties (aes for "aesthetic mapping"). Depending on the geom function, you can map variables to the x- and y-axis, to the shape of symbols, the line type of line graphs, to the fill color of symbols and more.  
> <br>
**Scales** can be customized with `scale_[]_[]` functions. These can be used to for example, change a color palette, set limits to the values included in the mapping, or log transform an axis.  
<br>
Modifications of the plot appearance (including background colors, thick labels, thick lines, axis labels, titles) can be achieved by combinations of several functions.  
<br>
A template code for a ggplot2 plot:  
`ggplot(data = <DATA>)+`  
  `geom_<FUNCTION>(mapping = aes(<MAPPING>))+`   
  `<SCALE_FUNCTION>+`
  `<THEME_FUNCTION>+`
  `<GGTITLE>`  
<br>

> **References**:  
> *[R for Datascience, chapter 3 Data visualization](https://r4ds.had.co.nz/data-visualisation.html)*  
> *[ggplot2 cheatsheet](https://github.com/rstudio/cheatsheets/raw/master/data-visualization-2.1.pdf)*  
 

### Exercise 3  

> Make a barplot showing the number of peaks per ChIP experiment.   
>  
> * All chip peaks are concatenated in the GRanges object `monocytes_all`.   

> **Exercise 3a**: Print a summary of `monocytes_all` using `show()` and identify the variable (column in metadata) that you can use for counting the number of peaks per chip (this will be the statistic used in the barplot).   

```{r monocytesall, exercise=TRUE, exercise.eval = F, error = T}
# print summary of object 'monocytes_all'


```

<div id="monocytesall-hint">
**Hint:** Use the `show()` function   
</div>

```{r monocytesall-solution, eval = F}
show(monocytes_all)
```

```{r monocytesall-check, eval = F}
grade_code(
  pass_if(~identical(.result, show(monocytes_all))))
```

```{r barplotvariable}
question("Which variable will you use for counting the number of peaks per ChIP-seq dataset?",
         answer("score"),
         answer("seqnames"),
         answer("peak"), 
         answer("chip", correct =T),
         allow_retry = T)
```

As GRanges objects are not compatible with ggplot2, you have to extract the metadata and write into a a data.frame  

### Exercise 3b  

> **Exercise 3b**: Complete the code below to extract the metadata from the object `monocytes_all`.  
> 
> * The function `as.data.frame()` is used to convert the resulting 'DFrame' object into a 'data.frame'.

```{r q_extractmetadata, exercise=T}
# extract metadata 
monocytes_metadata <- __(__)

# convert to a data.frame
monocytes_metadata <- as.data.frame(monocytes_metadata)
```

```{r q_extractmetadata-hint-1, eval = F}
monocytes_metadata <- mcols(monocytes_all)
```

```{r q_extractmetadata-solution, eval = F}
# extract metadata 
monocytes_metadata <- mcols(monocytes_all)

# convert to a data.frame
monocytes_metadata <- as.data.frame(monocytes_metadata)
```

```{r q_extractmetadata-code-check}
grade_code()
```

### Exercise 3c  

> **Exercise 3c**: Complete the code below to plot the number of peaks per ChIP in a bar chart using ggplot2 plotting.  

```{r prepare-barplotpeakcount} 
monocytes_metadata <- as.data.frame(mcols(monocytes_all))
```

```{r barplotpeakcount, exercise=TRUE, exercise.eval = F, error = T}
# plot the number of peaks per chip
__(data = __)+
  __(mapping=aes(y=__) )+
  scale_fill_colorblind()+
  theme_calc()+
  theme(legend.position = "none")+
  ggtitle("Number of peaks per ChIP, monocytes, chr19")
```

```{r barplotpeakcount-hint-1, eval = F}
ggplot(data = monocytes_metadata)
```

```{r barplotpeakcount-hint-2, eval = F}
geom_bar()
```

```{r barplotpeakcount-hint-3, eval = F}
ggplot(data = monocytes_metadata)+
  geom_bar(mapping=aes(y=chip))
```

```{r barplotpeakcount-solution, eval = F}
# plot the number of peaks per chip
ggplot(data = monocytes_metadata)+
  geom_bar(mapping=aes(y=chip))+
  theme_bw()+
  theme(legend.position = "none")+
  ggtitle("Number of peaks per ChIP, monocytes, chr19")
```

```{r q_barplotpeakcount-code-check}
grade_code()
```


*For which marks do you observe the most peaks? And for which the fewest number of peaks? Did they, in the genome browser, show a sharp peaks or broad domains of enrichment?*  
<br>
The range of peak counts can tell you whether your experiment and peak calling performed as expected. If they show extremely high or low values, it is likely something is off (something you should have already noted in the Genome Browser). It can also function as an extra control step to check that all peaks are indeed imported in R (and no strange things happened to your files along the way).   
<br>
In this case, the number of peaks is correct and you can continue with the analysis. 
<br>

### Exercise 4   

> Visualize the differences in peak size between the different marks in a boxplot.  
>  
> **Exercise 4a**: Finish the code below to add the 'peak_size' variable to `monocytes_metadata`. This new variable holds the peak widths.  

```{r prepare-addpeaksize, echo = F}
monocytes_metadata <- as.data.frame(mcols(monocytes_all))
```

```{r addpeaksize, exercise = T}
# add variable 'peak_size' to monocytes_metadata with the width of the peaks from monocytes_all 


```

<div id="addpeaksize-hint">
**Hint:** Use `monocytes_metadata$peak_size <- width(monocytes_all)`     
</div>

```{r addpeaksize-solution, eval = F}
monocytes_metadata$peak_size <- width(monocytes_all)
```

```{r addpeaksize-code-check, eval = F}
grade_code()
```

> **Exercise 4b**: Now plot the peak size distribution as a boxplot. Make sure to hit "Run Code" to show the plot in the console.     

```{r prepare-plotpeaksize}
monocytes_metadata <- as.data.frame(mcols(monocytes_all))
monocytes_metadata$peak_size <- width(monocytes_all)
```

```{r plotpeaksize, exercise=TRUE, exercise.eval = F, error = T}
# plot the peak_size per chip as boxplot
__(data = __)+
  __(mapping = __(x = __, y = __), fill = "gray", col = "black", bins = 100)+
  theme_bw()

```

<div id="plotpeaksize-hint">
**Hint:** Use `ggplot(data = monocytes_metadata)+`     
</div>

<div id="plotpeaksize-hint">
**Hint:** and `geom_boxplot(mapping = aes(x = chip, y = ...))`     
</div>

<div id="plotpeaksize-hint">
**Hint:** and y as `geom_boxplot(mapping = aes(x = chip, y = peak_size))`   
</div>

```{r plotpeaksize-solution, eval = F}
ggplot(data = monocytes_metadata)+
  geom_boxplot(mapping = aes(x = chip, y = peak_size), fill = "gray", col = "black")+
  theme_bw()
```

```{r plotpeaksize-code-check, eval = F}
grade_code()
```


> **Exercise 4c**: This doesn't look very informative, let's log10-transform the y-axis (again, hit "Run Code" to view the plot):  

```{r prepare-plotpeaksizelog}
monocytes_metadata <- as.data.frame(mcols(monocytes_all))
monocytes_metadata$peak_size <- width(monocytes_all)
```

```{r plotpeaksizelog, exercise=TRUE, exercise.eval = F, error = T}
# plot the peak_size per chip as boxplot, log10-transform the y-axis
__(data = __)+
  __(mapping = __(x = __, y = __), fill = "gray", col = "black")+
  theme_bw()+
  scale_y_continuous(trans = "__")
```

<div id="plotpeaksizelog-hint">
**Hint:** Specify "log10" for the `trans = ` parameter
</div>

```{r plotpeaksizelog-solution, eval = F}
ggplot(data = monocytes_metadata)+
  geom_boxplot(mapping = aes(x = chip, y = peak_size), fill = "gray", col = "black")+
  theme_bw()+
  scale_y_continuous(trans = "log10")
```

```{r plotpeaksizelog-code-check, eval = F}
grade_code()
```

As mentioned in section 2.2.1 and 2.2.6, different data types have different peak shapes. H3K4me3 and H3K4me1 ChIP-seq peaks are resp. **narrow**, and the **broad** peaks. 

```{r q_broad_vs_narrow}
quiz(caption = "Use the boxplot above.", 
     question("Which ChIP-seq dataset has/have broad peak shapes?",
         answer("H3K27ac"),
         answer("3K36me3", correct = T),
         answer("H3K27me3", correct = T),
         answer("H3K9me3", correct = T),
         allow_retry = TRUE),
     question("which have narrow peak shapes?",
         answer("H3K27ac", correct = T),
         answer("3K36me3"),
         answer("H3K27me3"),
         answer("H3K9me3"),
         allow_retry = TRUE,
         message = "You could have gotten this information from the [BLUEPRINT ChIP-Seq Analysis Pipeline description](http://dcc.blueprint-epigenome.eu/#/md/chip_seq_grch38) (remember: FAIR principles!). Under the section **Peak Calling** you can read that peak calling was performed with MACS2 and that the -broad flag was added depending on the mark in question." )
)
```
<br>

><details><summary>In summary, marks with narrow and with broad peaks:</summary>  
>  
> narrowPeak file | broadPeak file  
> :-- |:--  
H3K27ac | H3K27me3   
H3K4me3 | H3K36me3  
H3K9/14ac | H3K9me3  
H2A.Zac | H3K4me1  
</details>  

### 2.3.3 Plotting summarized values using dplyr & ggplot  

> In this section you will visualize the fraction that each ChIP-seq peak set covers in total on chromosome 19.  

To generate this plot, you have to calculate the total number of bps covered by each peak dataset. This can be achieved by summing all the width values.  
<br>
GRanges objects follow the **tidy data principle**: each row of a Ranges object corresponds to an interval, and each column will represent a variable about that interval, and generally each object will represent a single unit of an observation (like gene annotations). You can use **dplyr**-like functions from the `plyranges` package to manipulate these objects and use the **pipe** operator `%>%` to combine functions in a workflow.  

>#### **Refresher**: dplyr     
<details><summary>Common dplyr-functions and examples of their plyranges-relatives on GRanges objects</summary>   
>
Function | Purpose | Example on GRanges   
:--|:----|:----     
`select()` | subset variables (=columns) |  `select(my_gr, GC)`
`group_by()` | group data into rows with the same value for the specified variable. |  `my_gr %>% group_by(strand)`.  
`filter()` | subset observations (= rows) | `filter(my_gr, GC < 0.4)` or `my_gr %>% group_by(strand) %>%  filter(GC == max(GC))`.     
`summarize()` | Summarise variables, often per group | `group_by(my_gr, strand) %>% summarize(n = n(), gc = max(GC))`   
`mutate()` | Add a new varialbe | 
>
See also the [dplyr cheatsheet](https://rstudio.com/wp-content/uploads/2015/02/data-wrangling-cheatsheet.pdf)  
</details>  

### Exercise 5a  

> To plot the fraction of bps covered by each chip-seq peak dataset first calculate this fraction.  

> **Exercise 5a**: Finish the code below using `group_by()` and `summarize()` to calculate  
> 
> a. the total number of bps covered by each peak dataset and   
> b. as a fraction of total bps in chromosome 19 (=58617616 bp)   


```{r prepare-summarizewidth}
monocytes_metadata <- as.data.frame(mcols(monocytes_all))
monocytes_metadata$peak_size <- width(monocytes_all)
```

```{r summarizewidth, exercise=TRUE, exercise.eval = F, error = T}
# calculate the total number of bps covered by each chipseq dataset
# and as a fraction of total bps in chromosome 19
monocytes_metadata %>% 
  ___ %>%
  summarize(total_bps = ___ ,
            fraction_chr19 = ___ /58617616)
```

<div id="plotpeaksize-hint">
**Hint:** Use `monocytes_metadata %>% group_by(chip)`     
</div>

<div id="plotpeaksize-hint">
**Hint:** piped as `monocytes_metadata %>% group_by(chip) %>% summarize(total_bps = sum(peak_size))`     
</div>

<div id="plotpeaksize-hint">
**Hint:** and and the second summary variable defined as `monocytes_metadata %>% group_by(chip) %>% summarize(total_bps = sum(peak_size), fraction_chr19 = total_bps /58617616)`   
</div>

```{r summarizewidth-solution, eval = F}
monocytes_metadata %>% group_by(chip) %>% 
  summarize(total_bps = sum(peak_size),
            fraction_chr19 = total_bps /58617616)
```

```{r summarizewidth-code-check}
grade_code()
```

### Exercise 5b  

> **Exercise 5b**: To plot the `fraction_chr19` as bar chart, pipe the output of the code below into ggplot.  
> We instruct the bar function to plot the actual values with: `stat = "identity"`.  
> Hit "Run Code" to plot the plot in the console.  

```{r prepare-plotsummarizewidth, message = F}
monocytes_metadata <- as.data.frame(mcols(monocytes_all))
monocytes_metadata$peak_size <- width(monocytes_all)
```

```{r plotsummarizewidth, exercise=TRUE, exercise.eval = F, error = T}
# plot the fraction covered in a bar chart
monocytes_metadata %>%  group_by(chip) %>% 
  summarise(total_bps = sum(peak_size),  
            fraction_chr19 = sum(peak_size) /58617616) %>% 
  ggplot(.)+ 
  __(mapping = aes(x = __, y = __, fill = chip), stat = "identity")+
  theme_bw()+
  theme(legend.position = "none")
```

```{r plotsummarizewidth-hint-1, eval = F}
monocytes_metadata %>%  group_by(chip) %>% 
  summarise(total_bps = sum(peak_size),  
            fraction_chr19 = sum(peak_size) /58617616) %>% 
ggplot(.)+ 
  geom_bar(mapping = aes(x = chip ))
```

```{r plotsummarizewidth-hint-2, eval = F}
monocytes_metadata %>%  group_by(chip) %>% 
  summarise(total_bps = sum(peak_size),  
            fraction_chr19 = sum(peak_size) /58617616) %>% 
ggplot(.)+ 
  geom_bar(mapping = aes(x = chip ))
```

```{r plotsummarizewidth-solution, eval = F}
monocytes_metadata %>%  group_by(chip) %>% 
  summarise(total_bps = sum(peak_size),  
            fraction_chr19 = sum(peak_size) /58617616) %>% 
  ggplot(.)+ 
  geom_bar(mapping = aes(x = chip, y = fraction_chr19, fill = chip), stat = "identity")+
  theme_bw()+
  theme(legend.position = "none")
```

```{r plotsummarizewidth-code-check}
grade_code()
```

```{r mcsummarizedwidth}
question("Which type of peaks covers relativley more bp in chromosomes 9?",
         answer("narrow peak ChIP-seq peak datasets"),
         answer("broad domain ChIP-seq peak datasets", correct = T),
         allow_retry = T)
```

**These summary characteristics may suggest that H3K4me3 is relatively sparse (relative compared to the other marks) but that doesn't mean it isn't functional. Several analyses can inform us about the functionality of ChIP-seq marks. Let's start with their enrichment in promoters.**.
<br>

***
<br>


## 2.4 Overlap analysis   
### 2.4.1 Detect overlap with GenomicRanges       
To understand the function of a ChIPped histone mark or factor often want to know whether it is enriched in a particular genomic element. To find out, you need to calculate a quantitative summary of its genome-wide distribution across the different elements. By comparing this distribution to the genome-wide coverage of these different elements (null hypothesis), we can test for enrichment of the mark or factor.   
<br>
The genomic elements are defined in a reference file. You can use various reference files for this exercise: from transcript-oriented objects, to functional elements from a database like ENCODE, to a custom one.   
<br>
In the coming exercise, you will first quantify the overlap between H3K4me3 peaks and promoters and, subsequently, between H3K4me3 peaks and several different genomic elements simultaneously.   
<br>
To identify and count overlap between two sets of GRanges objects (e.g. H3K4me3 peaks and promoters) you can use function from the GenomicRanges package:  
<br>
**countOverlaps**: `countOverlaps(query, subject)` returns a integer vector with *the number of overlaps* for each element in the `query`.    
<br>
**subsetByOverlaps**: `subsetByOverlaps(query, subject)` extracts the elements in the `query` that overlap with at least one element in the subject.   
<br>
**findOverlaps**: `findOverlaps(query, subject)` returns a `Hits` object containing the index pairings for the overlapping elements.   

- The columns with indexes can be accessed through `queryHits(overlap_object)` and `subjectHits(overlap_object)`. These can subsequently be used to extract the corresponding peaks or regions from the objects that were used as inputs.    
- If a peak in one of the inputs overlaps with multiple peaks in the other, its index will appear multiple times in the output.  

<br>  

### 2.4.2 Overlap promoters vs h3k4me3 peaks   
To determine the overlap among H3K4me3 peaks and promoters, you first need to obtain promoter interval regions as GRanges object. We have done that for you with the following code that:  
  
1. Loads the `TxDb.Hsapiens.UCSC.hg38.knownGene` database package.  
2. Extract promoter regions with the `promoters()` function. This function outputs a GRanges object with intervals around the TSS.  
3. By default, `promoters()` works on transcripts but we are often more interested in promoters of genes therefore we add the `genes()` function.  

```{r retrieve_genes_txdb, exercise = T}
# load the txdb package which holds transcript-based gene models of hg38 genome  
library(TxDb.Hsapiens.UCSC.hg38.knownGene)
txdb <- TxDb.Hsapiens.UCSC.hg38.knownGene 
seqlevels(txdb) <- "chr19" # limits the database to chromosome 19

# extract promoter coordinates, filter for promoters 
promoters_chr19 <- unique(promoters(genes(txdb), upstream=2000, downstream=200)) 
```

### Exercise 6a 

> **Exercise 6a:** How many promoter are stored in `promoters_chr19`?     

```{r q6a_proms, exercise = T, exercise.lines = 3}
#  How many regions do you have are stored in `promoters_chr19`? 

```

<div id="q6a_proms-hint">
**Hint:** You may want to use the `length()` function.  
</div>

```{r q6a_proms-solution, eval = F}
length(promoters_chr19)
```

```{r q6a_proms-check}
grade_result(
 pass_if(~identical(.result, length(promoters_chr19)))
 )
```

### Exercise 6b 

> **Exercise 6b:** Use code to identify the width of the promoter windows. The output of this code should be one integer value.  

```{r q6b_proms, exercise = T, exercise.lines = 3}
# determine the promoter width

```

<div id="q6b_proms-hint">
**Hint:** first run a code to show a summary of the widths `summary(width(object))`  
</div>

<div id="q6b_proms-hint">
**Hint:** obtain the correct output format with `unique()`   
</div>

<div id="q6b_proms-hint">
**Hint:** `unique(width(promoters_chr19))`   
</div>

```{r q6b_proms-check}
grade_result(
 pass_if(~identical(.result, 2200))
 )
```


### Exercise 6c  

> We want to know how many of the promoters overlap with a H3K4me3 ChIP-seq peak and vice versa.    

> **Exercise 6c:** Use `findOverlaps()` to determine the overlap between `monocytes_h3k4me3` (query) and  `promoters_chr19` (subject).   
> 
> * You need to hit "Run Code" to be able to see the output of `show()`   


```{r q6c_overlap, exercise=TRUE, exercise.eval = F, error = T}
#exercise.setup = "prepare_exercise6"}
# Find overlap between monocytes_h3k4me3 peaks and promoters
overlap <- ___(query = ___, subject = ___)

# print the overlap output
show(___)
```

```{r q6c_overlap-solution, eval = F}
# Find overlap between monocytes_h3k4me3 peaks and promoters
overlap <- findOverlaps(query = monocytes_h3k4me3, subject = promoters_chr19)

# print the overlap output
show(overlap)
```

```{r q6c_overlap-code-check}
grade_code()
```
<br>
The Hits object "overlap" reports the indexes of `monocytes_h3k4me3` and `promoters_chr19` that overlap. If a peak or a promoter overlaps several times, each overlap will be reported in a new row.  
<br>

### Exercise 6d      

> **Exercise 6d:** What fraction of promoters overlaps with H3K4me3 peaks?  


```{r prepare_exercise6}
overlap <- findOverlaps(query = monocytes_h3k4me3, subject = promoters_chr19)
```

```{r q6d_overlap, exercise=TRUE, exercise.eval = F, error = T, exercise.setup = "prepare_exercise6"}
# Extract the unique promoter indexes
unique_promoters <- unique(___(overlap))

# How many of these unique promoters do you have?
unique_promoter_count <- ___(unique_promoters)

# Print the resulting number
unique_promoter_count

# What fraction of promoters is part of the overlap?  
unique_promoter_count/__(promoters)
```

<div id="q6d_overlap-hint">
Use `queryHits()` and `subjectHits()` to extract the indexes of overlapping peaks and promoters from `overlap`, respectively.  
</div>

<div id="q6d_overlap-hint">
Use `unique()` to minimize this output to unique peaks. 
</div>

<div id="q6d_overlap-hint">
Use `length()` to count the number of unique indexes of the query and the subject. 
</div>

```{r q6d_overlap-solution, eval = F}
# Extract the unique promoter indexes
unique_promoters <- unique(subjectHits(overlap))

# How many of these unique promoters do you have?
unique_promoters_count <- length(unique_promoters)

# Print the resulting number
unique_promoters_count

# What fraction of promoters is part of the overlap?  
unique_promoters_count/length(promoters_chr19)
```

```{r q6d_overlap_textbox, echo = FALSE}
question_text(
  "Enter the number of promoters that are part of the overlap:",
  answer("1723", message = "This is the total number of promoters. Find the ones part of the overlap with subjectHits() unique() and length()"),
  answer("1150", correct = TRUE),
  answer("1366", message = "This is the length of the overlap object, select the subject and reduce with unique."),
  answer("1087", message = "This is number of unique H3K4me3 peaks that are part of the overlap. Make sure you use subjectHits() and not queryHits() function."),
  allow_retry = TRUE)
```

```{r prepare_exercise6d}
overlap <- findOverlaps(query = monocytes_h3k4me3, subject = promoters_chr19)
unique_promoters <- unique(subjectHits(overlap))
unique_promoters_count <- length(unique_promoters)
```

```{r q6d_overlap_enterfraction, exercise = T, exerecise.lines = 2, exercise.setup = "prepare_exercise6d"}
# enter the fraction of promoters that are part of the overlap, round to 3 decimal places:  


```

```{r q6d_overlap_enterfraction-solution, eval = F}
round(unique_promoters_count/length(promoters_chr19),3)
```

```{r q6d_overlap_enterfraction-check}
grade_result(
  pass_if(~identical(.result,round(unique_promoters_count/length(promoters_chr19),3))))
```


### Exercise 6e

> **Exercise 6e:** And what fraction of the h3k4me3 peaks are part of the overlap?  

```{r q6e_overlap, exercise=TRUE, exercise.eval = F, error = T, exercise.setup = "prepare_exercise6"}
# Extract the unique peaks
unique_peaks <- unique(___(overlap))

# How many of these unique peaks do you have?
unique_peaks_count <- ___(unique_peaks)

# Print the resulting number
unique_peaks_count

# What fraction of peaks is part of the overlap?  
unique_peaks_count/__(__)
```

```{r q6e_overlap-solution, eval = F}
# Extract the unique peaks
unique_peaks <- unique(queryHits(overlap))

# How many of these unique peaks do you have?
unique_peaks_count <- length(unique_peaks)

# Print the resulting number
unique_peaks_count

# What fraction of peaks is part of the overlap?  
unique_peaks_count/length(monocytes_h3k4me3)
```

```{r q6e_overlap_textbox, echo = FALSE}
question_text(
  "Enter the number of H3K4me3 peaks that are part of the overlap:",
  answer("2838", message = "This is the total number of H3K4me3 peaks. Find the ones part of the overlap with subjectHits() unique() and length()"),
  answer("1087", correct = TRUE),
  answer("1366", message = "This is the length of the overlap object, select the subject and reduce the output with unique()."),
    answer("1150", message = "This is number of unique promoter regions that are part of the overlap. Make sure you use queryHits() and not the subjectHits() function."),
  allow_retry = TRUE)
```

```{r prepare_exercise6e}
overlap <- findOverlaps(query = monocytes_h3k4me3, subject = promoters_chr19)
unique_peaks <- unique(queryHits(overlap))
unique_peaks_count <- length(unique_peaks)
```

```{r q6e_overlap_enterfraction, exercise=TRUE, exercise.eval = F, error = T, exercise.setup =  "prepare_exercise6e"}
# enter the fraction of promoters that are part of the overlap, round to 3 decimal places:  


```

```{r q6e_overlap_enterfraction-solution, eval = F}
round(unique_peaks_count/length(monocytes_h3k4me3),3)
```

```{r q6e_overlap_enterfraction-check}
grade_result(
  pass_if(~identical(.result,round(unique_peaks_count/length(monocytes_h3k4me3),3))))
```

### Exercise 6f   
Plot the overlap among promoters and H3K4me3 peaks with the `plot.pairwise.venn()` function of the `VennDiagram` package.   

```{r prepare_exercise6f, echo = F}
overlap <- findOverlaps(query = monocytes_h3k4me3, subject = promoters_chr19)
# unique peaks 
unique_peaks <- unique(queryHits(overlap))
unique_peaks_count <- length(unique_peaks)
# unique promoters
unique_promoters <- unique(subjectHits(overlap))
unique_promoters_count <- length(unique_promoters)
# common
common_counts <- min(unique_promoters_count, unique_peaks_count)
```

We Use the minimum of the 'unique_peaks_count' and 'unique_promoters_count' as the number of 'common peaks' for our venn diagram.  
 
> **Exercise 6f:**   
> Complte the the `plot.pairwise.venn()` code below to visualize these counts in a venn diagram. The `plot.pairsiwe.venn()` function needs two "areas" that represent the total counts (overlapping and not-overlapping) of each element, and a count for the overlap.  
>
> * Hit "Run Code" to ensure that the venn diagram is plotted.  

```{r q6f_plotvenn, exercise = TRUE, exercise.eval= FALSE, error = T, exercise.setup = "prepare_exercise6f"}
# determine the number of common counts
common_counts <- min(unique_promoters_count, unique_peaks_count)

# call a new plotting area
grid.newpage()

# Plot the overlap in a venn diagram
draw.pairwise.venn( 
           area1=___, # total count for area 1
           area2=___, # total count for area 2
           cross.area=common_counts, # count for the overlap.
           category=c("H3K4me3", "Promoters"), 
           fill=c("red", "gray"), 
           cat.cex=1.2)
```

<div id="q6f_plotvenn-hint">
Use `length()` for area1 and area2
</div>


```{r q6f_plotvenn-solution, eval = F}
# determine the number of common counts
common_counts <- min(unique_promoters_count, unique_peaks_count)

# call a new plotting area
grid.newpage()

# Plot the overlap in a venn diagram
draw.pairwise.venn( 
   area1=length(monocytes_h3k4me3),
   area2=length(promoters_chr19), 
   cross.area=common_counts, 
   category=c("H3K4me3", "Promoters"), 
   fill=c("red", "gray"), 
   cat.cex=1.2)
```

```{r q7b_plotvenn-code-check}
grade_code()
```

**This looks like a considerable enrichment of H3K4me3 in promoter regions. Let's test this.**  

### 2.4.3 Testing for enrichment by overlap      

### Exercise 7: Is there a significant enrichment of h3k4me3 in promoters?
<br>
To answer this question, we compare the fraction of promoters with a H3K4me3 peak with the chromosome 19-wide fraction of promoters.
<br>
If H3K4me3 is not enriched at promoters, we would expect that the fraction of promoters with a H3K4me3 peak is in the same range as the fraction of promoters on chromosome 19. (This is our **null hypothesis of no enrichment**).
<br>
As some promoters may overlap we *reduced* `promoters` to non-overlapping intervals to calculate the chromosome 19-wide fraction of promoters.  We do that by running:
```{r fraction_promoters, echo = T, eval = F}
# obtain reduced promoter regions
promoters_chr19_reduced <- reduce(promoters_chr19)
```
Similarly to exercise 3 we then extract the metadata from this object and write it into a data.frame:
```{r metadata_prom, echo = T, eval = F}
# retrieve metadata and convert to dataframe
promoters_metadata <- as.data.frame(mcols(promoters_chr19_reduced))
```

### Exercise 7a

> **Exercise 7a**: Finish the code below to define the "promoter_size" variable and calculate the fraction of bps covered by `promoters_chr19_reduced` in chromosome 19 (chr19 =58617616 bp).

```{r q7a_promfraction-setup}
promoters_chr19_reduced <- reduce(promoters_chr19)
# convert to a data.frame
promoters_metadata <- as.data.frame(mcols(promoters_chr19_reduced))
```

```{r q7a_promfraction, exercise=TRUE, exercise.eval = F, error = T}
# define the variable 'promoter_size'
promoters_metadata$promoter_size <- ___(promoters_chr19_reduced)

# calculate the  fraction of bps covered 'promoters_chr19_reduced`
promoters_metadata %>%
  summarize(total_bps = ___ ,
            fraction_chr19 = ___ /58617616)
```

<div id="q7a_promfraction-hint">
**Hint:** Look back at exercise 5a.
</div>

```{r q7a_promfraction-solution, eval = F}
# define the variable 'promoter_size'
promoters_metadata$promoter_size <- width(promoters_chr19_reduced)

# calculate the  fraction of bps covered 'promoters_chr19_reduced`
promoters_metadata %>%
  summarize(total_bps = sum(promoter_size),
            fraction_chr19 = total_bps /58617616)
```

```{r q7a_promfraction_enterfraction, exercise = T}
# enter the fraction of promoters in chromosome 19, round of to 3 decimals


```

```{r q7a_promfraction_enterfraction-solution, echo = F}
round(0.06179487	,3)
```

```{r q7a_promfraction_enterfraction-check}
grade_result(
  pass_if(~identical(.result,round(0.06179487,3))))
```
<br>
Promoters make up 6% of this chromosome but appr. 66% of all promoters overlap with h3k4me3 peaks. This looks like an enrichment.

### Exercise 7b
You can test whether the observed fraction is indeed larger than expected with a binomial test, in r we can use the function `binom.test(x, p)` for this:

- `x` = vector with number of successes (= number of promoters with H3K4me3 peak) and number of failures (= number of promoters without H3K4me3 peak)
- `p` = expected probability of success, in this case the fraction of promoters in chromosome 19

<br>

> *Background*: The binomial test is run when an experiment has two possible outcomes (i.e. success/failure) and you have an idea about what the probability of success is. Success in this case is overlap and our expectation is that 20% of the cases show overlap. The test calculates the probability of getting a desired outcome with a specific sample size.

>**Exercise 7b**: Use a binomial test to test for enrichment of H3K4me3 in promoters.
>
> - Use the function `binom.test(x = c(successes, failures), p, alternative = "greater")` to call the test.
> - Use the information from exercise 6a and 6d to calculate the number of successes and failures.
> - Set `alternative = "greater"` because we test for *enrichment* and our alternative hypothesis is that the true probability is *larger* than the expected probability.

```{r prepare_q7b}
overlap <- findOverlaps(query = monocytes_h3k4me3, subject = promoters_chr19)
```

<br>
```{r q7b_binomtest, exercise=TRUE, exercise.eval = F, error = T}
# Use the indexes from the overlap object to identify promoters with H3K4me3 peak
promoters_with_h3k4me3 <- promoters_chr19[ unique(___(overlap)) ]

# Perform negative selection using the same indexes to find promoters without H3K4me3 peak
# the - sign removes the lines with the respective indexes from the original object
promoters_without_h3k4me3 <- promoters_chr19[ -unique(___(overlap)) ]

# Define the x for your test
# hint: use the function length()
x_test <- c(___, ___)

# Define the p for your test
# hint: exercise 7a
p_test <- __

# binomial test for enrichment of h3k4me3 peaks in promoters:
binomtest_result <- binom.test(x = x_test_prom_with_h3k4me3, p = p_test_prom_with_h3k4me3, alternative = "greater" )

# report the test output
binomtest_result
```

```{r q7b_binomtest-solution, eval = F}
# Use the indexes from the overlap object to identify promoters with H3K4me3 peak
promoters_with_h3k4me3 <- promoters_chr19[ unique(subjectHits(overlap)) ]

# Perform negative selection using the same indexes to find promoters without H3K4me3 peak
# the - sign removes the lines with the respective indexes from the original object
promoters_without_h3k4me3 <- promoters_chr19[ -unique(subjectHits(overlap)) ]

# Define the x for your test
# hint: use the function length()
x_test_prom_with_h3k4me3 <- c(length(promoters_with_h3k4me3), length(promoters_without_h3k4me3))

# Define the p for your test
# hint: exercise 7a
p_test_prom_with_h3k4me3 <- round(0.06179487,3)

# binomial test for enrichment of h3k4me3 peaks in promoters:
binomtest_result <- binom.test(x = x_test_prom_with_h3k4me3, p = p_test_prom_with_h3k4me3, alternative = "greater")

# report the test output
binomtest_result
```

### Exercise 7c
> **Exercise 7c** Based on this test, do you conclude that H3K4me3 is enriched in promoter regions?

```{r q7c_enrichmenttest}
question("Is H3K4me3 enriched promoters?",
         answer("*Yes*", correct = T, message = "The binomial test shows a p-value < 2.2e-16 and a confidence interval that excludes the expected 0.062 We therefore reject the H0 of no enrichment."),
         answer("*No*", message = "Incorrect. Look at the p-value of the previous test.")
)
```

### 2.4.4 What are the genes with a H3K4me3 peak in their promoter?  
To find the genes associated with the promoters, we use the indexes from the overlap object to subset the original `promoters_chr19` object and extract the corresponding geneIDs.

### Exercise 8  

> **Exercise 8**: How many different genes are associated withthe H3K4me3-overlapping promoters? 

```{r prepare_exercise8, echo = F}
overlap <- findOverlaps(query = monocytes_h3k4me3, subject = promoters_chr19)
```

```{r q8a_genesids_overlap, exercise=TRUE, exercise.eval = F, error = T, exercise.setup = "prepare_exercise8"}
# the objects `overlap` and `promoters_chr19` are loaded for you 
# subset for promoters that are part of the overlap with H3K4me3 in monocytes 
promoters_with_h3k4me3  <- 

# convert the metadata of this object into a dataframe
# count the length of unique geneIDs using summarize() and ndistinct()
as.data.frame(mcols(promoters_with_h3k4me3)) %>% 
  ___(unique_geneids = ___ )  
```

<div id="q8a_genesids_overlap-hint">
Use the 'gene_id' column in the object `promoters_chr19`   
</div>

<div id="q8a_genesids_overlap-hint">
Use `subjectHits(overlap)` to subset `promoters_chr19`    
</div>

<div id="q8a_genesids_overlap-hint">
`promoters_chr19[subjectHits(overlap)]`    
</div>

<div id="q8a_genesids_overlap-hint">
To get a vector of geneIDs, you need to 
</div>

<div id="q8a_genesids_overlap-hint">
`length(unique(promoters_with_h3k4me3[,1]))`   
</div>

```{r q8a_genesids_overlap-solution, eval = F}
# the objects `overlap` and `promoters_chr19` are loaded for you 
# subset for promoters that are part of the overlap with H3K4me3 in monocytes 
promoters_with_h3k4me3 <- promoters_chr19[subjectHits(overlap)]

# count the length of unique geneIDs in this object using summarize() and ndistinct()
as.data.frame(mcols(promoters_with_h3k4me3)) %>% 
  summarise(unique_geneids = n_distinct(gene_id)) 
```

```{r q8a_genesids_overlap-check}
grade_result(
  pass_if(~identical(.result,1150)))
```

## 2.5 Comparing gene expression   
### 2.5.1 Genes with H3K4me3 at their promoter  
Can we observe a difference in gene expression between genes with and without H3K4me3 in their promoters?

### Exericse 9: Plot the expression of genes with and without H3K4me3 in their promoters  
We obtained the gene expression quantification from the BLUEPRINT data portal [ref](http://dcc.blueprint-epigenome.eu/#/experiments/ERX157053) and filtered it for genes present on chromosome 19. These data are loaded in object `quantification_chr19` for you.  
<br>

> **Exercise 9a**: first use dplyr::mutate(NEW_VARIABLE = ....) to define a new column `h3k4me3_promoter` which groups genes based on their overlap (TRUE/FALSE) with H3K4me3 peaks.  

<br>
```{r prepare_exercise9a, echo = F}
overlap <- findOverlaps(query = monocytes_h3k4me3, subject = promoters_chr19)
promoters_with_h3k4me3 <- promoters_chr19[subjectHits(overlap)]
```

```{r q9a_geneexpression, exercise=TRUE, exercise.eval = F, error = T, exercise.setup = "prepare_exercise9a"}
# The object quantification_chr19 holds the RNA-seq quantification for genes on chromosome 19. 
# View the structure to identify the variable you will use in mutate. 
str(quantification_chr19)

# Define a new variable that identifies if a gene is present in the overlap or not.  
quantification_chr19_new <- quantification_chr19 %>% 
  ___(h3k4me3_promoter = ___ %in% mcols(promoters_with_h3k4me3)[,1] )
```

```{r q9a_geneexpression-solution, eval =F}
# The object quantification_chr19 holds the RNA-seq quantification for genes on chromosome 19. 
# View the structure to identify the variable you will use in mutate. 
str(quantification_chr19)

# Define a new variable that identifies if a gene is present in the overlap or not.  
quantification_chr19_new <- quantification_chr19 %>% 
  mutate(h3k4me3_promoter = entrezgene_id %in% promoters_with_h3k4me3$gene_id )
```

```{r q9a_geneexpression-code-check}
grade_code()
```

> **Exercise 9b**: Use ggplot() and geom_boxplot() to visualize the FPKM (y-axis) per gene group.  

```{r prepare_exercise9b, echo = F}
overlap <- findOverlaps(query = monocytes_h3k4me3, subject = promoters_chr19)
promoters_with_h3k4me3 <- promoters_chr19[subjectHits(overlap)]
quantification_chr19_new <- quantification_chr19 %>% 
  mutate(h3k4me3_promoter = entrezgene_id %in% promoters_with_h3k4me3$gene_id )
```

```{r q9b_plot_genequant, exercise=TRUE, exercise.eval = F, error = T, exercise.setup = "prepare_exercise9b"}
# Use ggplot() and geom_boxplot() to visualize the FPKM per gene group in the dataset `quantification_chr19_new`   
ggplot(___)+
  geom_boxplot(mapping = aes(x = __, y = __))+
  theme_calc()
```

```{r q9b_plot_genequant-solution, eval = F}
ggplot(quantification_chr19_new)+
  geom_boxplot(mapping = aes(x = h3k4me3_promoter, y = FPKM))+
  theme_calc()
```

```{r q9b_plot_genequant-code-check}
grade_code()
```

This doesn't look very informative. RNA-seq quantification has a high range. We must therefore log-transoform the y-axis. 

> **Exercise 9c**: Repeat the plot of 9b but lgo10-transform the y-axis.  
>
> - add 1 to all FPKM values to prevent log10(0).  
> - also add another 'layer', visualizing the same data but as jitter (points).  

```{r q9c_plot_genequant, exercise=TRUE, exercise.eval = F, error = T, exercise.setup = "prepare_exercise9b"}
# Repeat the plot of 9b but lgo10-transform the y-axis    
ggplot(quantification_chr19_new)+
  geom_boxplot(mapping = aes(x = h3k4me3_promoter, y = (FPKM+1)))+
  geom_jitter(mapping = aes(x = h3k4me3_promoter, y = (FPKM+1)), point= 0.5, alpha = 0.6)+ # alpha makes the points transparant, ensuring that you still see the violin shape.  
  theme_calc()+
  scale_y_continuous(trans = "___")
```

```{r q9c_plot_genequant-solution, eval = F}
# Repeat the plot of 9b but lgo10-transform the y-axis    
ggplot(quantification_chr19_new)+
  geom_boxplot(mapping = aes(x = h3k4me3_promoter, y = (FPKM+1)))+
  geom_jitter(mapping = aes(x = h3k4me3_promoter, y = (FPKM+1)), point= 0.5, alpha = 0.8)+ 
  theme_calc()+
  scale_y_continuous(trans = "log10")
```

```{r q9c_plot_genequant-code-check}
grade_code()
```

There is a clear difference in gene expression between the two gene sets with presence of H3K4me3 being associated with higher expression. H3K27me3 can also be deposited at promoters. What trend do you observe there?  

### 2.5.2 Genes with H3K27me3 @ promoter   
### Exercise 10: Plot the expression of genes with and without H3K27me3 at their promoters  

> **Exercise 10**: To do this, you must  
>  
1. identify overlap between H3K27me3 peaks and promoters  
2. obtain promoters that are reported in the overlap  
3. define a new variable called "h3k27me3_promoter" in `quantification_chr19` that groups genes based on the overlap of their promoter with H3K27me3 peak. Instead of writing this into a new object, you can pipe the output data.frame immediately into step 4.    
4. plot the FPKM per gene set. 
* Use the previous exercises if you need hints to code these steps. 


```{r q10_H3K27me3_prom, exercise =  T , exercise.eval = F, error = T}
# 1. identify the overlap between monocytes_h3k27me3 and promoters_chr19
overlap_h3k27me3_promoters <- ___(query = ___, subject = ___)

# 2. obtain promoters that are reported in the overlap  
promoters_with_h3k27me3 <- promoters_chr19[___(overlap_h3k27me3_promoters)]

# 3. define a "h3k27me3_promoter" in `quantification_chr19`, pipe the output in your ggplot command  
quantification_chr19 %>%  
  ___(h3k27me3_promoter = ____) %>% 
  ggplot(.) +
  geom_boxplot(mapping = aes(x = ___, y = (FPKM+1)))+
  geom_jitter(mapping = aes(x = ___, y = (FPKM+1)), alpha = 0.8, size = 0.5)+
  theme_calc()+
  scale_y_continuous(trans = "log10")
```


```{r q10_H3K27me3_prom-solution, eval =F}
# 1. identify the overlap between monocytes_h3k27me3 and promoters_chr19
overlap_h3k27me3_promoters <- findOverlaps(query = monocytes_h3k27me3, subject = promoters_chr19)

# 2. obtain promoters that are reported in the overlap  
promoters_with_h3k27me3 <- promoters_chr19[subjectHits(overlap_h3k27me3_promoters)]

# 3. define a "h3k27me3_promoter" in `quantification_chr19`, pipe the output in your ggplot command  
quantification_chr19 %>%  
  mutate(h3k27me3_promoter = entrezgene_id %in% promoters_with_h3k27me3$gene_id) %>% 
  ggplot(.) +
  geom_boxplot(mapping = aes(x = h3k27me3_promoter, y = (FPKM+1)))+
  geom_jitter(mapping = aes(x = h3k27me3_promoter, y = (FPKM+1)), alpha = 0.8, size = 0.5)+
  theme_calc()+
  scale_y_continuous(trans = "log10")
```

```{r q10_H3K27me3_prom-code-check}
grade_code()
```

This time the trend is the other way around; presence of H3K27me3 is associated with lower gene expression. H3K4me3 and H3K27me3 can also co-occur. What gene expression level is associate with promoters that have both marks?  

### 2.5.3 Bivalent promoters   
This time you will make 4 gene sets, depending on the overlap of promoters with H3K4me3 and/or H3K27me3 peaks. 

### Exercise 11  

> **Exercise 11a**: To do this, you must define a new "promoter_mark" in `quantification_chr19` that groups genes based in 4 groups of promoters based on them being or not being reported in 'promoters_with_only_h3k4me3' or 'promoters_with_only_h3k27me3'. For this, we use `case_when()` within `mutate()` and define new category names based on the conditions as shown below.       

```{r prepare_exercise11, echo = F}
overlap_h3k4me3_promoters <- findOverlaps(query = monocytes_h3k4me3, subject = promoters_chr19)
promoters_with_h3k4me3 <- promoters_chr19[subjectHits(overlap_h3k4me3_promoters)]

overlap_h3k27me3_promoters <- findOverlaps(query = monocytes_h3k27me3, subject = promoters_chr19)
promoters_with_h3k27me3 <- promoters_chr19[subjectHits(overlap_h3k27me3_promoters)]
```

```{r q11a_bivalent_prom, exercise =  T , exercise.eval = F, error = T, exercise.setup = "prepare_exercise11"}
# 1. define the variable "promoter_mark" in `quantification_chr19`
# define the level name between the ""-signs.  
quantification_chr19_new <- quantification_chr19 %>% mutate(
  promoter_mark = case_when(
    entrezgene_id %in% promoters_with_h3k4me3$gene_id & entrezgene_id %in% promoters_with_h3k27me3$gene_id ~ "____",  # promoters with both marks
    entrezgene_id %in% promoters_with_h3k4me3$gene_id & !(entrezgene_id %in% promoters_with_h3k27me3$gene_id) ~ "____", # promoters with only h3k4me3
    !(entrezgene_id %in% promoters_with_h3k4me3$gene_id) & entrezgene_id %in% promoters_with_h3k27me3$gene_id ~ "____", # and so forth... 
    !(entrezgene_id %in% promoters_with_h3k4me3$gene_id | entrezgene_id %in% promoters_with_h3k27me3$gene_id) ~ "____") ) 

# print a table with the number of promoters per promoter mark

```

<div id="q11_bivalent_prom-hint">
An example of level names would be:  
`quantification_chr19 %>%  mutate(promoter_mark = case_when(`
    `entrezgene_id %in% promoters_with_h3k4me3$gene_id & entrezgene_id %in% promoters_with_h3k27me3$gene_id ~ "bivalent",` 
    `entrezgene_id %in% promoters_with_h3k4me3$gene_id & !(entrezgene_id %in% promoters_with_h3k27me3$gene_id) ~ "H3K4me3_only", `
    `!(entrezgene_id %in% promoters_with_h3k4me3$gene_id) & entrezgene_id %in% promoters_with_h3k27me3$gene_id ~ "H3K27me3_only",` 
    `!(entrezgene_id %in% promoters_with_h3k4me3$gene_id | entrezgene_id %in% promoters_with_h3k27me3$gene_id) ~ "no_mark") )`
</div>

```{r q11a_bivalent_prom-solution, eval = F}
quantification_chr19_new <- quantification_chr19 %>%  mutate(promoter_mark = case_when(
    entrezgene_id %in% promoters_with_h3k4me3$gene_id & entrezgene_id %in% promoters_with_h3k27me3$gene_id ~ "bivalent", 
    entrezgene_id %in% promoters_with_h3k4me3$gene_id & !(entrezgene_id %in% promoters_with_h3k27me3$gene_id) ~ "H3K4me3_only", 
    !(entrezgene_id %in% promoters_with_h3k4me3$gene_id) & entrezgene_id %in% promoters_with_h3k27me3$gene_id ~ "H3K27me3_only", 
    !(entrezgene_id %in% promoters_with_h3k4me3$gene_id | entrezgene_id %in% promoters_with_h3k27me3$gene_id) ~ "no_mark") ) 

table(quantification_chr19_new$promoter_mark)

```

```{r q11a_bivalent_prom-check}
grade_result(
  pass_if(~identical(.result,table(quantification_chr19_new$promoter_mark)))
)
```


> **Exercise 11b**: plot the gene expression as log10(FPKM+1) in a boxplot with extra geom_jittter, per group of genes.  

```{r prepare_exercise11b, echo = F}
overlap_h3k4me3_promoters <- findOverlaps(query = monocytes_h3k4me3, subject = promoters_chr19)
promoters_with_h3k4me3 <- promoters_chr19[subjectHits(overlap_h3k4me3_promoters)]

overlap_h3k27me3_promoters <- findOverlaps(query = monocytes_h3k27me3, subject = promoters_chr19)
promoters_with_h3k27me3 <- promoters_chr19[subjectHits(overlap_h3k27me3_promoters)]

quantification_chr19_new <- quantification_chr19 %>%  mutate(promoter_mark = case_when(
    entrezgene_id %in% promoters_with_h3k4me3$gene_id & entrezgene_id %in% promoters_with_h3k27me3$gene_id ~ "bivalent", 
    entrezgene_id %in% promoters_with_h3k4me3$gene_id & !(entrezgene_id %in% promoters_with_h3k27me3$gene_id) ~ "H3K4me3_only", 
    !(entrezgene_id %in% promoters_with_h3k4me3$gene_id) & entrezgene_id %in% promoters_with_h3k27me3$gene_id ~ "H3K27me3_only", 
    !(entrezgene_id %in% promoters_with_h3k4me3$gene_id | entrezgene_id %in% promoters_with_h3k27me3$gene_id) ~ "no_mark") ) 
```


```{r q11b_bivalent_prom, exercise =  T , exercise.eval = F, error = T, exercise.setup = "prepare_exercise11b"}
# pipe the output into your ggplot
  ggplot(___) +
    geom_boxplot(mapping = aes(x = ___, y = (FKPM+1)))+
    geom_jitter(mapping = aes(x = ___, y = (FKPM+1)), alpha = 0.8, size = 0.5)+
    theme_calc()+
    scale_y_continuous(trans = "log10")
```

<div id="q11b_bivalent_prom-hint">
`ggplot(quantification_chr19_new)+ geom_boxplot(mapping = aes(x = promoter_mark, y = (FPKM+1)))+`
    `geom_jitter(mapping = aes(x = promoter_mark, y = (FPKM+1)), alpha = 0.8, size = 0.5)+`
    `theme_calc()+`
    `scale_y_continuous(trans = "log10")``
</div>

```{r q11b_bivalent_prom-solution, eval =F}
ggplot(quantification_chr19_new) +
    geom_boxplot(mapping = aes(x = promoter_mark, y = (FPKM+1)))+
    geom_jitter(mapping = aes(x = promoter_mark, y = (FPKM+1)), alpha = 0.8, size = 0.5)+
    theme_calc()+
    scale_y_continuous(trans = "log10")
```

```{r q11_bivalent_mc}
question("Which of the following statments are correct?",
         answer("There are more promoters with only H3K4me3 than with H3K4me3 combined with another mark.", message = "We have to evaluate more marks to determine whether H3K4me3 is indeed found more often by itself at a promoter than combined with another mark."),
         answer("H3K27me3 is not doing anything as the associated genes have the same expression level as genes without marks."),
         answer("Gene expression is associated with differnt histone mark (combinations).", correct = T),
         answer("H3K4me3 defines promoters.", message = "H3K4me3 does not define promoters. Promoters are defined as the regions near the TSS."),
         answer("There are more promoters with chromatin markings than without, in chromosome 19 of monocytes.", correct = T),
         answer("H3K4me3 is more often found at promoters than H3K27me3, in chromosome 19 of monocytes.", correct = T),
         allow_retry = T,
         random_answer_order = T
         )

```


**So far we have only looked at individual overlaps but chromatin states are actually defined by the combinatorial occupancy of different marks and factors. Let's analyse the chromosome19-wide overlap of the marks in our dataset and identify the occuring combinations.**  

## 2.6 Summarizing overlap among marks   
You will make an UpSet plot, which is an elegant alternative to the VennDiagram, that visualizes the size of each detected combination. Combination in our case refers to overlap. This **size** of the combined set can be defined in various ways, depending on how you compare your regions. In the `ComplexHeatmap` package they define three 'modes' for it: distinct, intersect and union. You will work with the results from the setting 'distinct'. Below is a visual summary of these modes.   

![](images/wk2_4_complexheatmap_modes.png)/  
[Ref](https://jokergoo.github.io/ComplexHeatmap-reference/book/upset-plot.html#input-data): *Zuguang Gu, ComplexHeatmap Complete Reference, last revised on 2020-10-27*  
<br>
<br>
The `make_comb_mat([list_of_grangesobjects])` is used to detect overlapping marks:   
```{r makecombexample, eval = F, echo = T}
monocytes_combinationmatrix <- make_comb_mat(monocytes_list) # takes a while! 
```
`make_comb_mat()` returns a "combination matrix" object. Special functions can be used to view its information and filter the reported combinations:   

- set_name() prints the names of the objects from the original list  
- set_size() prints the total number of bps covered by the intervals in each object  
- comb_size() reports the total number of bps covered by the respective combination of intervals.  
    + The combinations are encoded in a binary fashion of 0s and 1s.  
    + A 1 means the mark in question is present in the combination.  
    + These 0s and 1s follow the order given by set_name().   


### Exericse 12a    

> **Exercise 12a**: Examine the `monocytes_combinationmatrix` object with the obove mentioned functions.     

```{r prepare_exercise12a, echo = F}
# import correct output 
### rdsfile <- system.file("extdata", "week2", "prepared_rds", "blueprint_monocyte_chr19_complexheatmap_combmatrix.rds", package = "learnr.proto")
rdsfile <- paste("data", "prepared_rds", "blueprint_monocyte_chr19_complexheatmap_combmatrix.rds", sep = "/")
monocytes_combinationmatrix <- readRDS(rdsfile)
```

```{r q12a_upsetrplot_mark_overlap, exercise = TRUE, error = T, exercise.eval = F, exercise.setup =  "prepare_exercise12a"}
# What are the names of the sets?
__(monocytes_combinationmatrix)

# What are the sizes of the sets? (in million bps)
___(monocytes_combinationmatrix)/1e-6

# What are the sizes of the detected set combinations? (in million bps)
___(monocytes_combinationmatrix)/1e-6
```

```{r q12a_upsetrplot_mark_overlap-solution, eval = F}
# What are the names of the sets?
set_name(monocytes_combinationmatrix)

# What are the sizes of the sets? (in million bps)
set_size(monocytes_combinationmatrix)/1e-6

# What are the sizes of the detected set combinations?  
comb_size(monocytes_combinationmatrix)/1e-6
```

```{r q12_bivalent_mc}
quiz(caption ="", 
     question("Which mark covers most bps?",
              answer("H3K4me3"),
              answer("H3K36me3"),
              answer("H3K27me3"),
              answer("H3K9me3", correct = T, message = "other marks have more peaks, but this marks spans broad domains."),
              allow_retry = T),
     question("And which combination of marks apparantly covers most bps?",
              answer("H3K4me3 and H3K27me3"),
              answer("H3K9me3 and H3K27me3"),
              answer("H3K9me3 and H3K4me3"),
              answer("H3K9me3 and H3K36me3", correct = T, message = "These two marks generally cover large windows in the genome. Their co-occupancy has been associated weakly transcribed genes. A large domain of H3K36me3 H3K9me3 bivalent chromatin is found at the cluster of zinc finger genes on chromosome 19 we saw in week1. (ref: 
                     Mauser et al. 2017 Epigentics and Chromatin)"),
              allow_retry = T,
              random_answer_order = T)
     )
```


The combination H3K4me1 and H3K4me3 is the 3rd most widespread combination. This combination might well represent signal due to some cross-reactivity of the antibody (-ies) or co-chipping of these regions because of enhancer-promoter pairing. It would help if we could also look at the signal besides its presence/absence. Will do so in week3.  

### Exericse 12b  

> **Exercise 12a**: Filter `monocytes_combinationmatrix` for combinations of 100kb and more. And plot the filtered combination matrix with `UpSet()`.   


```{r q12b_upsetrplot, exercise = T, exercise.setup = "prepare_exercise12a", exercise.eval = F, error = T}
# filter monocytes_combinationmatrix for combinations covering at least 100kb
monocytes_combmatrix_filt <- monocytes_combinationmatrix[comb_size(monocytes_combinationmatrix) > ___ ]

# plot the result in an upset plot:  
___(monocytes_combmatrix_filt)


```

```{r q12b_upsetrplot-solution, eval = F}
# filter monocytes_combinationmatrix for combinations covering at least 100kb
monocytes_combmatrix_filt <- monocytes_combinationmatrix[comb_size(monocytes_combinationmatrix) > 100000]

# plot the result in an upset plot:  
UpSet(monocytes_combmatrix_filt)
```

```{r q12b_upsetrplot-code-check}
grade_code()
```

### Exericse 12c  

> **Exercise 12c**: Run the code below, which changes the labels at the top and rightside.   

```{r q12c_niceupset, exercise = T, exercise.setup = "prepare_exercise12a"}
# filter monocytes_combinationmatrix for combinations covering at least 100kb
monocytes_combmatrix_filt <- monocytes_combinationmatrix[comb_size(monocytes_combinationmatrix) > 100000]

# plot with nice axis labels
UpSet(monocytes_combmatrix_filt, top_annotation = upset_top_annotation(
		monocytes_combmatrix_filt,
		axis_param = list(at = c(0, 2e6,4e6,6e6, 8e6, 10e6),
			labels = c("0Mb", "2Mb","4Mb", "6Mb", "8Mb", "10Mb")),
		height = unit(2, "cm")
	),
	right_annotation = upset_right_annotation(
		monocytes_combmatrix_filt,
		axis_param = list(at = c(0,4e6,8e6, 1.2e7),
			labels = c("0Mb", "4Mb", "8Mb", "12Mb"),
			labels_rot = 0),
		width = unit(2, "cm")
	))
```

This makes the plot a bit better readable. You can also add additional plots to the top and bottom of this graph.  

### Exericse 12d  

> **Exercise 12d**: Run the code below, calculating the distance to the nearest TSS and plotting those values in the 

First we filter the original combination matrix and then extract all the matrices into a GRanges object. The `lapply()` command returns a list of these GRanges objects.  

```{r add_distance_to_tss, echo = T}
# filter monocytes_combinationmatrix for combinations covering at least 200kb
monocytes_combmatrix_filt <- monocytes_combinationmatrix[comb_size(monocytes_combinationmatrix) > 100000]

# extract different combinations they are turned into a list of GRanges objects  
monocytes_combmatrix_sets <- lapply(comb_name(monocytes_combmatrix_filt), function(nm) extract_comb(monocytes_combmatrix_filt, nm))
```

```{r q12d_add_distance_to_tss, exercise = T, fig.height=5, exercise.setup = "add_distance_to_tss"}
# define variable "dist_to_tss" in each GRanges object
	# use distancetoNearest from genomicranges package
  # extract the metadat from the output and take column 1 with the distance in bp from that object
monocytes_combmatrix_sets <- lapply(monocytes_combmatrix_sets, function(gr) {
	gr$dist_to_tss <- mcols(distanceToNearest(gr, tss_chr19))[,1]
	gr$nearby_tss_geneid <- tss_chr19[subjectHits(distanceToNearest(gr, tss_chr19))]$gene_id
	gr
})


# plot the UpSet plot with nice top and right-side annotation. 
# add boxplots at the bottom represting the distance to the tss.  
UpSet(monocytes_combmatrix_filt, 
      top_annotation = upset_top_annotation(
        monocytes_combmatrix_filt,
        axis_param = list(at = c(0, 2e6,4e6,6e6, 8e6, 10e6),
                          labels = c("0Mb", "2Mb","4Mb", "6Mb", "8Mb", "10Mb")),
        height = unit(4, "cm")
        ),
      right_annotation = upset_right_annotation(
        monocytes_combmatrix_filt,
        axis_param = list(at = c(0,4e6,8e6, 1.2e7),
                          labels = c("0Mb", "4Mb", "8Mb", "12Mb"),
                          labels_rot = 0),
        width = unit(4, "cm")
        ),
      bottom_annotation = HeatmapAnnotation(
        dist_to_tss = anno_boxplot(lapply(monocytes_combmatrix_sets, function(gr) gr$dist_to_tss), outline = FALSE),
        label = anno_text(LETTERS[1:length(monocytes_combmatrix_sets)], rot = 0),
        annotation_name_side = "left",
        height = unit(4, "cm")
	))
```

```{r q12e_bivalent_mc}
 question("Use the plot above, which statements do you agree with?",
         answer("There are more promoters with only H3K4me3 than with H3K4me3 combined with another mark.", message = "Based on these marks, we would say no. Compare with the output from question 11. Taking more marks along let us see that H3K4me3 only is sparse."),
         answer("Column I (letters at the bottom), marks promoters", correct = T),
         answer("H3K4me3 defines promoters.", message = "H3K4me3 does not define promoters. Promoters are defined as the regions near the TSS."),
         answer("Column M marks the second-most widespread combination and could be promoters or enhancers.", correct = T),
         answer("Column D, with H3K36me3 only marking, likely represents silences genes"),
         allow_retry = T,
         random_answer_order = T
         )
```


## So far,... 
So far we have examined peak locations, overlap with features and the association between the absence or presence of a mark and gene expression. This binary approach is a very simplistic model and we have been ignoring the signal strength altogether. As in, *Are high and lowly expressed genes associated with equal levels of H3K4me3 at their promoter or not? And what about H3K36me3 across the gene body?*. The measure of strength is also taken along in the definition of different chromatin states, see for example this paper on the [chromHMM tool](https://doi.org/10.1038/nmeth.1906). Positional and signal information are combined to learn these states. Next week you will examine the relationship between signal strength and gene expression. 

But what also should be said, you have come very far already. Very well done, keep up the good work!!    

```{js print2pdf1, context="server"}
$(document).on('shiny:inputchanged', function(event) {
  if (event.name === 'print2pdf') {
    window.print();
  }
});
```

```{r print2pdf2}
actionButton("print2pdf", "Print page", style="opacity: .7; color: #000;")
```
